Modelling collective behavior in groups of mice housed under semi-naturalistic conditions

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #2 (Public review):

The authors have constructively responded to previous referee comments and I believe that the manuscript is a useful addition to the literature. I particularly appreciate the quantitative approach to social behavior, but have two cautionary comments.

(1) Conceptually it is important to further justify why this particular maximum entropy model is appropriate. Maximum entropy models have been applied across a dizzying array of biological systems, including genes, neurons, the immune system, as well as animal behavior, so would seem quite beneficial to explain the particular benefits here, for mouse social behavior as coarse-grained through the eco-hab chamber occupancy. This would be an excellent chance to amplify what the models can offer for biological understanding, particularly in the realm of social behavior

We thank the reviewer for this comment. Maximum entropy models, along with other statistical inference methods that learn interaction patterns from simultaneously-measured degrees of freedom, help distinguish various types of interactions, e.g. direct vs. indirect interactions among animals, individual preference to food vs. social interaction with pairs. As research on social behavior expands from focusing on pairs of animals to studying groups in (semi-)naturalistic environments, maximum entropy models serve as a crucial link between high-throughput data and the need to identify and distinguish interaction rules. Specifically, among all possible maximum entropy models, the pairwise maximum entropy model is one of the simplest that can describe interactions among individuals, which serves as an excellent starting point to understand collective and social behavior in animals.

Although the Eco-HAB setup currently records spatially coarse-grained data, it still provides more spatial information compared to the traditional three-chamber tests used to assess sociability for rodents. By showing that the maximum entropy model can effectively analyze Eco-HAB data, we hope to highlight its potential in research of social behavior in animals.

To amplify what the models can offer for biological understanding particularly in the realm of social behavior, We have updated the Introduction to add a more logical structure to the need of using maximum entropy models to identify interactions among mice. Additionally, we updated the first paragraph of the Discussion to make it specific that it is the use of maximum entropy models that identifies interaction patterns from the high-throughput data. Finally, we have also added in the Discussion (line 422-425) arguments supporting the specific use of pairwise maximum entropy models to study social behaviors.

(2) Maximum entropy models of even intermediate size systems involve a large number of parameters. The authors are transparent about that limitation here, but I still worry that the conclusion of the sufficiency of pairwise interactions is simply not general, and this may also relate to the differences from previous work. If, as the authors suggest in the discussion, this difference is one of a choice of variables, then that point could be emphasized. The suggestion of a follow up study with a smaller number of mice is excellent.

We thank the reviewer for raising the issue and agree that the caveat of how general pairwise interactions can describe social behavior of animals needs to be discussed. We have added a sentence in the Discussion to point out this important caveat. “More generally, this discrepancy when looking at different choices of variables raises the issue that when studying social behavior of animals in a group, it is important to test and compare interaction models with different complexity (e.g. pairwise or with higher-order interactions).” We have also toned down our conclusion to limit our results of pairwise interactions describing mice co-localization patterns to the data collected in Eco-HAB (also see Reviewer 3 Major Point 2).

Reviewer #3 (Public review):

Summary:

Chen et al. present a thorough statistical analysis of social interactions, more precisely, co-occupying the same chamber in the Eco-HAB measurement system. They also test the effect of manipulating the prelimbic cortex by using TIMP-1 that inhibits the MMP-9 matrix metalloproteinase. They conclude that altering neural plasticity in the prelimbic cortex does not eliminate social interactions, but it strongly impacts social information transmission.

Strengths:

The quantitative approach to analyzing social interactions is laudable and the study is interesting. It demonstrates that the Eco-HAB can be used for high throughput, standardized and automated tests of the effects of brain manipulations on social structure in large groups of mice.

Weaknesses:

A demonstration of TIMP-1 impairing neural plasticity specifically in the prelimbic cortex of the treated animals would greatly strengthen the biological conclusions. The Eco-HAB provides coarser spatial information compared to some other approaches, which may influence the conclusions.

Recommendations for the authors:

Reviewer #3 (Recommendations for the authors):

Major points

(1) Do the Authors have evidence that TIMP-1 was effective, as well as specific to the prelimbic cortex?

We refer to the literature for the effectiveness and specificity of TIMP-1 to the prelimbic cortex.

Specifically, the study by Okulski et al. (Biol. Psychiatry 2007) provides clear evidence that TIMP1 plays a role in synaptic plasticity in the prefrontal cortex. They showed that TIMP-1 is induced in the medial prefrontal cortex (mPFC) following stimulation that triggers late long-term potentiation (LTP), a key model of synaptic plasticity. Overexpression of TIMP-1 in the mPFC blocked the activity of matrix metalloproteinases (MMPs) and prevented the induction of late LTP in vivo. Similar effects were observed with pharmacological inhibition of MMP-9 in vitro, reinforcing the idea that TIMP-1 regulates extracellular proteolysis as part of the plasticity mechanism in the prefrontal cortex. These findings confirm that TIMP-1 is both effective and active in this specific brain region.

Further evidence comes from Puścian et al. (Mol. Psychiatry 2022), who used TIMP-1-loaded nanoparticles to influence neuronal plasticity in the amygdala. They found that TIMP-1 affected MMP expression, LTP, and dendritic morphology, showing its impact on synaptic modifications. More directly relevant, Winiarski et al. (Sci. Adv. 2025) demonstrated that injecting TIMP-1-loaded nanoparticles into the prelimbic cortex altered responses to social stimuli, further supporting the idea that TIMP-1 has region-specific effects on behavioral processes.

We have also updated the main text (page 8, 1st paragraph of “Effect of impairing neuronal plasticity in the PL on subterritory preferences and sociability”) of the manuscript to include the above references.

(2) The Authors seem to suggest that one main reason for the different results compared to Shemesh et al. 2013 was the coarseness of the Eco-HAB data. In this case, I think this conclusion should be toned down because of this significant caveat.

We thank the reviewer for pointing this out, and agree that this caveat and difference should be emphasized. To tone down the conclusion, we have

(1) added details about the Eco-HAB (it being coarse-grained, etc.) in the abstract to tone down the conclusion.

(2) added to the results summary in the Discussion (top of page 12) that the results are “within in the setup of the semi-naturalistic Eco-HAB experiments”

(3) added to the Discussion (page 13) that the different results compared to Shemesh et al 2013 means that general studies of social behavior need to compare models with different levels of complexity (e.g. pairwise vs. higher-order interactions). (Also see Reviewer 2 Comment 2.)

Minor points

(1) Please explain what is measured in Fig. 1C (what is on the y axis?).

Figure 1C shows the activity of the mice as measured by the rate of transitions, i.e. the number of times the mice switch boxes during each hour of the day, averaged over all N = 15 mice and T = 10 days (cohort M1). The error bars represent variability of activities across individuals or across days. For mouse-to-mouse variability (blue), we first compute for each mouse its number of transitions averaged over the same hour for all 10 days, then we compute its standard deviation across all 15 mice and plot it as error bars. For day-to-day variability (orange), we first compute for each day the number of transitions for each hour averaged over all mice, then compute its standard deviation across all 10 days as the errorbar. We have added the detailed explanation in the caption of Figure 1C.

(2) In Fig. 3, it would be better to present the control group also in the main figure instead of the supplementary.

We have merged Figure 3 and Figure 3 Supplementary 1 to present the control group also in the main figure.

(3) In Fig. 3 and corresponding supplements, there seems to be a large difference between males and females. I think this would deserve some more discussion.

While not being the main focus of this paper, we agree with the reviewer that the difference between male and female is important and deserves attention in the discussion and also future study. Thus we have added a paragraph in the Discussion (line 394-399, bottom of page 12).

Changes from the earlier version to the revised version

Abstract

Earlier version: In social behavior research, the focus often remains on animal dyads, limiting the understanding of complex interactions. Recent trends favor naturalistic setups, offering unique insights into intricate social behaviors. Social behavior stems from chance, individual preferences, and group dynamics, necessitating high-resolution quantitative measurements and statistical modeling. This study leverages the Eco-HAB system, an automated experimental setup which employs RFID tracking to observe naturally formed mouse cohorts in a controlled yet naturalistic setting, and uses statistical inference models to decipher rules governing the collective dynamics of groups of 10-15 individuals. Applying maximum entropy models unveils social rules in mouse hordes, quantifying sociability through pairwise interactions within groups, exploring how social structure evolves, the impact of individual versus social preferences, and the effects of considering interaction structures among three animals instead of two. Reproducing co-localization patterns of individual mice reveals stability over time, with the inferred interaction strength capturing social structure. By separating interactions from individual preferences, the study demonstrates that altering neuronal plasticity in the prelimbic cortex – the brain structure crucial for sociability – does not eliminate social interactions, but makes the transmission of social information between mice more challenging. The study demonstrates how the joint probability distribution of the mice positions can be used to quantify sociability.

Revised version: In social behavior research, the focus often remains on animal dyads, limiting the understanding of complex interactions. Recent trends favor naturalistic setups, offering unique insights into intricate social behaviors. Social behavior stems from chance, individual preferences, and group dynamics, necessitating high-resolution quantitative measurements and statistical modeling. This study leverages the Eco-HAB system, an automated experimental setup which employs RFID tracking to observe naturally formed mouse cohorts in a controlled yet naturalistic setting, and uses statistical inference models to decipher rules governing the collective dynamics of groups of 10-15 individuals. Applying maximum entropy models on the coarse-grained co-localization patterns of mice unveils social rules in mouse hordes, quantifying sociability through pairwise interactions within groups, the impact of individual versus social preferences, and the effects of considering interaction structures among three animals instead of two. Reproducing co-localization patterns of individual mice reveals stability over time, with the statistics of the inferred interaction strength capturing social structure. By separating interactions from individual preferences, the study demonstrates that altering neuronal plasticity in the prelimbic cortex – the brain structure crucial for sociability – does not eliminate signatures of social interactions, but makes the transmission of social information between mice more challenging. The study demonstrates how the joint probability distribution of the mice positions can be used to quantify sociability.

Introduction

Social behavior is fundamental for numerous animal species, encompassing human societies. From the dynamic spectacle of Mexican waves in a football stadium to the intricate waggle dance of bees, the diverse manifestations of social interaction raise a pivotal question: How do these social behaviors come to fruition, and what roles do individuals play in their emergence?

In recent decades, the exploration of social behavior has predominantly centered around studying animal dyads in controlled laboratory conditions. However, these experimental paradigms inherently impose limitations on investigating intricate social behaviors that often involve more than two interacting individuals. Studies on social interactions frequently employ tests with brief observation periods, during which animals are evaluated in novel environments, accompanied by the presence of an experimenter that induces stress, influencing social behavior [1–5]. Recently, there has been a notable shift towards conducting experiments in natural settings, involving animal groups such as flocks of birds [6] or swarms of midges [7]. Additionally, semi-naturalistic environments, exemplified by fish in tanks [8–10], marching locusts in arenas [11], flocks of sheep [12], and hordes of rodents [13, 14], are increasingly being utilized. These approaches present unique opportunities for the comprehensive quantification of complex social interactions and sociability.

Mice stand out as a valuable model system for delving into the complexities of social behavior, given their intricate manifestation of various social behaviors. They tend to form cohesive groups, showcasing both amicable and agonistic behaviors. Depending on the environmental context, mice demonstrate territoriality and dynamic social hierarchies [15]. Communication among mice is extensive, primarily mediated through odors, allowing them to convey emotional states such as stress, fear, and preferences in food [16, 17]. Additionally, mice exhibit prosocial behaviors, actively assisting distressed fellow mice in need [18]. Decades of research have extensively explored social interactions between pairs of mice, while the study of mouse groups has only recently become feasible with advancements in high-throughput technologies, particularly radiofrequency identification (RFID) [14, 19]. The Eco-HAB system, utilized in this study, leverages RFID tracking to observe naturally-formed cohorts of mice in a controlled yet naturalistic environment, enabling longitudinal experiments on sociability with minimal human interference [14].

Removed in the revised version: Social behavior arises from chance, individual preferences, group structure, and the transmission of preferences and interactions among group members. To unravel these elements and understand the establishment of social networks and hierarchies, we need high-resolution quantitative measurements of behavior over extended periods and statistical modeling to capture natural variability. In this study, we integrate Eco-HAB recordings with statistical inference to construct models of collective behavior, focusing on the statistics of system states to identify interaction structures within the group.

Removed in the revised version: The prefrontal cortex (PFC) plays a crucial role in processing social information, understanding others’ emotions, maintaining social hierarchy, and transmitting information about food safety in both rodents and humans [20, 21]. Neuronal activity of the PFC is correlated with proximity to conspecifics, and studies in mice reveal distinct PFC responses to social and non-social olfactory stimuli [22, 23]. The PFC integrates existing knowledge with new information about self and others, demonstrating dynamic neuronal plasticity [24]. In cognitive tasks involving the PFC and subcortical areas, neuronal connectivity refines more rapidly in the former, highlighting its adaptability to changing environments [25, 26]. Tissue inhibitors of metalloproteinases (TIMPs), particularly TIMP-1, influence synaptic plasticity by inhibiting matrix metalloproteinases (MMPs), especially MMP-9 [27–31]. TIMP-1 is involved in long-term potentiation (LTP), a crucial process for cellular-level memory formation [32, 33]. This sustained release impedes the updating of neuronal connectivity in the prelimbic part of the PFC (PL), crucial for maintaining social structure [34–36]. Our study employs nanoparticles for gradual TIMP-1 release over several days [37] to impact plasticity in the PL on the changes in group behavior.

Earlier version: Maximum entropy models successfully explain social rules governing collective behavior in bird flocks and mouse hordes [38, 39]. These models help distinguish observed correlations, like the clustering of mice in a specific location, from direct interactions or individual preferences. [39] pioneered the use of these models in studying mouse group behavior, revealing the significance of higher-order interactions in co-localization patterns. While [39] utilized video tracking of groups of four mice, our Eco-HAB setup employs RFID technology for tracking groups of 10-15 mice, providing more compact data with longer recording times but lower spatial resolution. Our focus is on quantifying sociability in mouse hordes through pairwise interactions within groups, ensuring statistical power. We explore whether interactions between pairs can explain collective behavior and examine how social structure evolves over time. We analyze the effects of individual versus social preferences and investigate the impact of considering three animals instead of two. Using a data analysis approach based on wild type C57BL6/J male and female mice, we discuss social structure and sociability changes in mice with temporary prefrontal cortex plasticity modification.

Revised version: Social behavior arises from a combination of chance, individual preferences, group structure, and the transmission of preferences and interactions among group members. To unravel these elements and understand the establishment of social networks and hierarchies, we need not only high-resolution quantitative measurements of behavior over extended periods, but also statistical modeling to construct interaction models of collective behavior. One particular statistical method that has been successfully applied to identify interaction models in a diverse range of biological networks are maximum entropy models. Among many examples, maximum entropy models have successfully explain social rules governing collective behavior in bird flocks and mouse hordes [20, 21]. These models help distinguish observed correlations, like the clustering of mice in a specific location, from direct interactions or individual preferences. [21] pioneered the use of these models in studying mouse group behavior, revealing the significance of higher-order interactions in colocalization patterns. While [21] utilized video tracking of groups of four mice, our Eco-HAB setup employs RFID technology for tracking groups of 10-15 mice, providing more compact data with longer recording times but lower spatial resolution. In this study, we integrate Eco-HAB recordings with statistical inference to construct models of collective behavior, focusing on the statistics of system states to identify interaction structures within the group. Our focus is on quantifying sociability in mouse hordes through the inferred interactions within groups, ensuring statistical power. We explore whether interactions between pairs can explain collective behavior and examine how social structure evolves over time. We analyze the effects of individual versus social preferences and investigate the impact of considering three animals instead of two. Using a data analysis approach based on wild type C57BL/6J male and female mice, we discuss social structure and sociability changes in mice with temporary prefrontal cortex plasticity modification.

Added in the revised version: Furthermore, we explore whether interventions of brain regions that are crucial for processing social information change the interaction patterns among individuals in the Eco-HAB. The prefrontal cortex (PFC) plays a crucial role in processing social information, understanding others’ emotions, maintaining social hierarchy, and transmitting information about food safety in both rodents and humans [22, 23]. Neuronal activity of the PFC is correlated with proximity to conspecifics, and studies in mice reveal distinct PFC responses to social and non-social olfactory stimuli [24, 25]. The PFC integrates existing knowledge with new information about self and others, demonstrating dynamic neuronal plasticity [26]. In cognitive tasks involving the PFC and subcortical areas, neuronal connectivity refines more rapidly in the former, highlighting its adaptability to changing environments [27, 28]. Tissue inhibitors of metalloproteinases (TIMPs), particularly TIMP-1, influence synaptic plasticity by inhibiting matrix metalloproteinases (MMPs), especially MMP-9 [29–33]. TIMP-1 is involved in long-term potentiation (LTP), a crucial process for cellular-level memory formation [34, 35]. This sustained release impedes the updating of neuronal connectivity in the prelimbic part of the PFC (PL), crucial for maintaining social structure [36–38]. Our study employs nanoparticles for gradual TIMP-1 release over several days [39] to impact plasticity in the PL on the changes in group behavior.

Results

Recording of mice location in naturalistic environment

Earlier version: Eco-HAB is an automated, ethologically-relevant experimental apparatus that tracts voluntary behavior in group-housed mice [14]. Constructed to simulate notable characteristics of natural murine environment, it consists of four connected large compartments, two of which contain food and water (Fig. 1A). Cohorts of 10 to 15 mice are introduced into the Eco-HAB, where they behave freely while their locations are tracked over time. The details of used mouse strains and cohorts’ compositions can be found in Material and Methods. The compartments are connected with tube-shaped corridors resembling underground tunnels, on whose ends there are 125kHz antennas recording every time a mouse crosses with an accuracy of over 20Hz. Each mouse is tagged with a unique RFID tag. The mice are recorded for 10 days with alternating 12-hour-long light-dark phases that simulate the day-night cycle.

Revised version: Eco-HAB is an automated, ethologically-relevant experimental apparatus that tracts voluntary behavior in group-housed mice [14]. Constructed to simulate notable characteristics of natural murine environment, it consists of four connected large compartments, two of which contain food and water (Fig. 1A). Cohorts of 10 to 15 mice are introduced into the Eco-HAB, where they behave freely while their locations are tracked over time. The details of used mouse strains and cohorts’ compositions can be found in. The compartments are connected with tube-shaped corridors resembling underground tunnels, on whose ends there are 125kHz antennas recording every time a mouse crosses with an accuracy of over 20Hz. Each mouse is tagged with a unique RFID tag. The mice are recorded for 10 days with alternating 12-hour-long light-dark phases that simulate the day-night cycle.

Mice were tested in Eco-HAB, a system for automated, ecologically-relevant assessment of voluntary behavior in groups of mice. Animals were tested for 10 days.

Earlier version: (A) Schematic of the Eco-HAB system, where four compartments are connected with tunnels. Food and water are available ad libitum in compartments 2 and 4. (B) Time series of the location of 15 mice over 10 days, as aligned to the daylight cycle. (C) The activity of the mice is affected by the circadian clock. Error bars represent standard deviation across all mice (mouse-mouse variability, in blue) or across all days for the mean activity level for all mice (day-day variability, in orange). The two curves are slightly shifted horizontally for clearer visualization. We focus the following analysis on the data collected during the first half of the dark phase, between 13:00 and 19:00 (shaded region).

Revised version: (A) Schematic of the Eco-HAB system, where four compartments are connected with tunnels. Food and water are available ad libitum in compartments 2 and 4. (B) Time series of the location of 15 mice over 10 days, as aligned to the daylight cycle. (C) Circadian clock affect the activity of the mice, measured by the number of transitions in each hour averaged over the 15 mice and the 10 days. Error bars represent standard deviation across all mice (mouse-mouse variability, in blue) or across all days for the mean activity level for all mice (day-day variability, in orange). The two curves are slightly shifted horizontally for clearer visualization. We focus the following analysis on the data collected during the first half of the dark phase, between 13:00 and 19:00 (shaded region).

The location of each mouse at each time is reconstructed using the time stamps, reducing the data to a discrete time series, σt at time t = 1, 2, …, T, with possible values of the locations σt = 1, 2, 3, 4 corresponding to the four compartments. The time resolution for the discretization is set to 2 seconds. As shown by the color-coded location traces in Fig. 1B, the majority of mice are often found in the same compartment, especially in the non-active light phases: this corresponds to the ethological behavior – mice tend to sleep in a pile to keep each other warm (see [40]). This suggests that the behavior depends on latent variables, i.a. the circadian clock.

To ensure the relative consistency of the analyzed data, we used the observed rate of transitions to measure the activity of each mouse, and choose an analysis window of 6 hours covering the first half of the dark phase (13:00 - 19:00), which corresponds to the most active time on each day for the entire duration of the experiment. The variability of activity across individuals is larger than the day-to-day variability for a single mouse, suggesting that the level of locomotor activity is a well-defined individual characteristic (Fig. 1C).

Pairwise interaction model explains the statistics of social behavior

We first establish a quantification of sociability by building probabilistic interaction models for groups of mice. Following previous work [14], we use the in-cohort sociability, which measures the excess probability of two mice being found in the same compartment compared to the case where they are independent. Mathematically, in-cohort sociability is defined as: where and fij(r, r′) are respectively the empirical frequencies of finding a mouse i in compartment r, and a pair of mice (i, j) in compartments r, r′.

As schematically explained in Fig. 2A, in-cohort sociability is due to pairwise interactions between each pair of mice, and modifies how likely they are to be found in each compartment with respect to the mice’s innate preference for that compartment. However, considering the presence of more than two animals, in-cohort sociability is not an effective measure of social structure of the group: two animals with zero attraction to one another can still be found to have a high in-cohort sociability, if a third animal has a strong social bond with both of them, since they all will be spending time with one another.

Mice in Eco-HAB interact pairwisely.

Earlier version: (A, B) The schematics showing pairwise interactions: two mice are more likely to be found in the same compartment than the sum of their individual preference implies (A); also, the probability for three mice being in the same compartment can be predicted from the pairwise interactions (B). (C) From pairwise correlation Cij, defined as the probability for mouse i and mouse j being in the same compartment (subtracted by the prediction of the independent model), and the probability for mouse i to be found in compartment r (subtracted by the model where each mouse spends equal amount of time in each of the four compartments), mir − 1/4, pairwise maximum entropy model learns the interaction strength between a pair of mice, Jij, and the local field , which gives the tendency for each mouse in each compartment. The data shown is the aggregated 5-day data from day 1 to day 5 of the C57 males (cohort M1). The pairwise maximum entropy model can predict higher order statistical structures of the data (schematics in panel B), such as the probability for triplets of mice being in the same compartment (subtracted by the prediction of the independent model, mathematically fijk(r, r, r) − fi(r)fj(r)fj(r)) (D), and the probability of K mice being found in the same compartment (E). (F) Conditional log-likelihood of mouse locations, predicted by the pairwise model (l(1,2)), the independent model (l(1)), and the null model assuming no compartment preference or interactions (l(0), the yellow line), for each mice (cohort M1 before TIMP-1 treatment, N = 15).

Revised version: (A, B) The schematics showing pairwise interactions: two mice are more likely to be found in the same compartment than the sum of their individual preference implies (A); also, the probability for three mice being in the same compartment can be predicted from the pairwise interactions (B). (C) From pairwise correlation Cij, defined as the probability for mouse i and mouse j being in the same compartment (subtracted by the prediction of the independent model), and the probability for mouse i to be found in compartment r (subtracted by the model where each mouse spends equal amount of time in each of the four compartments), mir − 1/4, pairwise maximum entropy model learns the interaction strength between a pair of mice, Jij, and the local field hi, which gives the tendency for each mouse in each compartment. The data shown is the aggregated 5-day data from day 1 to day 5 of the C57BL/6J males (cohort M1). The pairwise maximum entropy model can predict higher order statistical structures of the data (schematics in panel B), such as the probability for triplets of mice being in the same compartment (subtracted by the prediction of the independent model, mathematically fijk(r, r, r) − fi(r)fj(r)fj(r)) (D), and the probability of K mice being found in the same compartment (E). (F) Conditional log-likelihood of mouse locations, predicted by the pairwise model (l(1,2)), the independent model (l(1)), and the null model assuming no compartment preference or interactions (l(0), the yellow line), for each mice (cohort M1 before TIMP-1 treatment, N = 15).

Since measurements of location preference, and the incohort sociability, together with the dynamical observables such as the rate of activity, are stable over time (SI Fig. S1) it invites a quantitative modeling of the joint-probability distribution of the co-localization of mice.

Earlier version: To distinguish social structure interactions from the effective correlations that define in-cohort sociability, we build a maximum entropy model with pairwise interactions. This approach constrains the joint probability distribution of all the possible co-localization patterns of all mice to reproduce the empirical occupation frequencies and the in-cohort sociability , while otherwise remaining as random as possible [39, 41, 42]. With these assumptions, the joint probability distribution of the mice co-localization patterns can be written as where hir is the individual preference of mouse i to be in compartment r, and Jij is the interaction between mouse i and mouse j. The set of parameters ({hir, Jij}) is learned through gradient descent (see SI Fig. S2 and Materials and Methods for details). The interactions Jij may be positive or negative. We see that although the structure of the interactions Jij follows that of incohort sociability Cij, they are not identical (SI Fig. S3). Likewise, individual mice preferences hir are not equal to the occupation probability mir (Fig. 2C). Thus, this approach allows us to distinguish direct interactions from indirect ones.

Revised version: To distinguish social structure interactions from the effective correlations that define in-cohort sociability, we build a maximum entropy model with pairwise interactions. This approach constrains the joint probability distribution of all the possible co-localization patterns of all mice to reproduce the empirical occupation frequencies and the in-cohort sociability , while otherwise remaining as random as possible [21, 41, 42]. With these assumptions, the joint probability distribution of the mice co-localization patterns can be written as where hir is the individual preference of mouse i to be in compartment r, and Jij is the interaction between mouse i and mouse j. The set of parameters ({hir, Jij}) is learned through gradient descent (see SI Fig. S2 and Materials and Methods for details). The interactions Jij may be positive or negative. We see that although the structure of the interactions Jij follows that of incohort sociability Cij, they are not identical (SI Fig. S3).

Added in the revised version: Likewise, individual mice preferences hir are not equal to the occupation probability mir (Fig. 2C). Thus, this approach allows us to distinguish direct interactions from indirect ones.

Earlier version: To validate the model, we tested that it is able to predict higher order features of the data, such as the probability of a specific combination of triplets of mice being in the same compartment (Fig. 2BD), and the probability of observing K mice in the same compartment (Fig. 2E), with the overestimation at large K possibly due to the limit of finite data. Although the model assumes the strength of interaction does not depend on which compartment the mice are in, our minimal model can predict probability of K mice in certain compartments (SI Fig. S4, compartments 1 and 3). We call in-state probability the distribution of box occupancy of each mouse given the position of all other mice (see Material and Methods). The model prediction for in-state probabilities match the observed one, showing that the model gives an unbiased estimate of individual mouse positions SI Fig. S5). Moreover, models with triplet interactions show signs of overfitting under cross-validation, which is mitigated when the triplet interactions are suppressed close to zero using L2 regularization (see Material and Methods and SI Fig. S6). These results show that pairwise interaction among mice are sufficient to assess the observed collective behavior.

Revised version: To validate the model, we tested that it is able to predict higher order features of the data, such as the probability of a specific combination of triplets of mice being in the same compartment (Fig. 2BD), and the probability of observing K mice in the same compartment (Fig. 2E), with the overestimation at large K possibly due to the limit of finite data. Although the model assumes the strength of interaction does not depend on which compartment the mice are in, our minimal model can predict probability of K mice in certain compartments (SI Fig. S4, compartments 1 and 3). We call in-state probability the distribution of box occupancy of each mouse given the position of all other mice (see Materials and Methods). The model prediction for in-state probabilities match the observed one, showing that the model gives an unbiased estimate of individual mouse positions (SI Fig. S5). Moreover, models with triplet interactions show signs of overfitting under cross-validation, which is mitigated when the triplet interactions are suppressed close to zero using L2 regularization (see and SI Fig. S6). These results show that pairwise interaction among mice are sufficient to assess the observed collective behavior.

Choosing timescales for analysis

Earlier version: To construct an interaction model based on the steady state distribution, we first need to consider the proper timescales for which we average over the observables, i.e. mean and correlation of the mice co-localization patterns. If the timescale is too short, then the error of estimation may be large. More severely, the system may not have enough time to equilibrate, and the time average will not result in the steady state distribution. On the other hand, if the timescale is too long, we loose biologically meaningful information about the temporal evolution of the system, such as the adaptation of the mice in a new environment and evolution of the social interaction strength.

Revised version: To construct an interaction model based on the steady state distribution, we first need to consider the proper timescales for which we average over the observables, i.e. mean and correlation of the mice co-localization patterns.

Added in the revised version: If the timescale is too short, then the error of estimation may be large. More severely, the system may not have enough time to equilibrate, and the time average will not result in the steady state distribution. On the other hand, if the timescale is too long, we loose biologically meaningful information about the temporal evolution of the system, such as the adaptation of the mice in a new environment and evolution of the social interaction strength.

Earlier version: To identify the proper timescale, we systematically conduct cross-validation for pairwise maximum entropy models constructed using K days of data, where K = 1, 2, 3, 4, 5, 10, and each day of the data contains the 6 hours when the mice are most active. The data is then separated into training sets that consist of 5-hours of the data from each day and test sets that consist of 1-hour of the data each day. Pairwise maximum entropy models with L2 regularization with strength βJ imposed on the pairwise interactions are learned from the training set, and the training- and the test-set likelihoods are computed. As shown in SI Fig. S7, the test-set likelihood decreases as the regularization strength increases for cumulate data with number of days K≥ 4, indicating that the pairwise model generalizes well. We choose K = 5 for subsequent analysis, as it does not overfit, and it gives us temporal information about how the interaction structures may change over the 10-day experiment.

Revised version: To identify the proper timescale, we systematically conduct cross-validation for pairwise maximum entropy models constructed using K days of data, where K = 1, 2, 3, 4, 5, 10, and each day of the data contains the 6 hours when the mice are most active. The data is then separated into training sets that consist of 5-hours of the data from each day and test sets that consist of 1-hour of the data each day. Pairwise maximum entropy models with L2 regularization with strength βJ imposed on the pairwise interactions are learned from the training set, and the trainingand the test-set likelihoods are computed. As shown in SI Fig. S7, the test-set likelihood decreases as the regularization strength increases for cumulate data with number of days K ≥ 4, indicating that the pairwise model generalizes well. We choose K = 5 for subsequent analysis, as it does not overfit, and it gives us temporal information about how the interaction structures may change over the 10-day experiment.

Stability of sociability over time

Earlier version: The data-driven model and its inferred parameters allow us to explore various aspects of social behavior. As the models are built using accumulated data from 5 days in the 10-day experiment, we first assess the temporal consistency of the chamber preference hir and the inferred interaction parameters Jij of the four cohorts of C57BL6/J male mice (see Table 1), M1 (N = 15), M2 (N = 13), M3 (N = 10), M4 (N = 12, before BSA injection).

Revised version: The data-driven model and its inferred parameters allow us to explore various aspects of social behavior. As the models are built using accumulated data from 5 days in the 10-day experiment, we first assess the temporal consistency of the chamber preference hir and the inferred interaction parameters Jij of the four cohorts of C57BL/6J male mice (see Table 1), M1 (N = 15), M2 (N = 13), M3 (N = 10), M4 (N = 12, before BSA injection).

As shown in SI Fig. S8, for all four cohorts, the box preference tendency is consistent, with a Pearson’s correlation coefficient between models learned from the first and the last five days being 0.6 ± 0.1. For cohort M1, M3 and M4, the distribution of the individual chamber preferences is consistent over time, with the two-sample F-test for equal variance being non-significant, and the two-sample t-test for equal mean being either non-significant or with a p-value of 0.04 for cohort M1. For cohort M2, both the mean and the variance of the chamber preferences significantly changed between the first and the last 5 days of the experiment. Nonetheless, the Pearson’s correlation between the two data segments remain large at p = 0.57, indicating a consistency over time for the same cohort. Different cohorts exhibit different distributions of chamber preferences.

For the inferred interaction parameters Jij, the distribution is consistent between the first 5 days and the last 5 days of the data. Specifically, for all four cohorts, the standard deviations of the interactions do not change between the first 5 days and the last 5 days of the data, as shown by two-sample F-test for equal variance. The mean of the inferred interactions does not change for cohort M1 and M3, however for cohort M2 and M4, two-sample t-test for equal mean returns a p = 0.0034 and p = 0.0073 respectively for the interactions. Notably, across from the four different cohorts, two-sample F-test with Bonferroni correction shows that the variance of all eight 5-day modeling are not significantly different. In contrast, cross-cohort comparison between male (cohort M1 before drug injection) and female (cohort F1 before drug injection) shows significantly different variance (p < 0.01 between the last 5 days of F1 and the first 5 days or the last 5 days of M1) and mean (p < 0.001) of the inferred interaction strength, which shows that these measures of sociability can be used to distinguish strains or genders. Nonetheless, the strengths of the individual interactions in the specific pairs of mice i and j, Jij, vary more notably, as given by Pearson’s correlation coefficient 0.015 ± 0.165 (see SI Fig. S8C). This implies that the maximum entropy model does not infer a social structure that is stable over time.

Quantifying the influence of social versus individual preferences

Further, we ask how important social interactions are for determining mice behavior, by measuring how much the data can be explained by the individual preferences for specific spaces within the territory vs. the interactions with other mice. Mice are social animals, yet they perform many behaviors based on their individual moment-to-moment needs, and it is unclear a priori how much the social interactions influence mice behavior in comparison to their individual preference.

For each mouse i, we consider three nested models with increasing descriptive power: first, the null model assuming each mouse has the same probability of being found in each compartment, P (0)(σi) = 1/4; second, the independent model that assumes no interactions among mice, and the probability of finding each mouse in each compartment is solely determined by their individual preferences, P (1)(σi) = fi(σi); third, the inferred pairwise interaction model based on voluntarily spending time with other mice considered, P (2)(σi |{σj≠i}), using Eq. 2.

We then quantified how well each model explains the data by comparing the mean log-likelihoods of finding a mouse in a given compartment, conditioned on the location of all other mice. As shown in Fig. 2F, including information on pairwise interactions increases the log-likelihood of the data by as much as including information on individual compartment preferences, as shown by the similar values of the probability ratios P (2)/P (1) and P (1)/P (0). The likelihood is consistent between the first five days and the last five days of the experiment, but exhibits variability across different cohorts of animals within the same strain (SI Fig. S9).

Earlier version: Another possible measure of sociability is the mutual information between a single mouse’s location within the territory and the location of the rest of the cohort, which tells us how accurately the position of a single mouse can be predicted if the positions of all other mice are known (see details in Material and Methods). The possible values of the mutual information is between 0 and 2 bits, where 0 bits means no predictability, and 2 bits means perfect predictability. In our Eco-HAB data, the average mutual information for each mouse is 0.0323± 0.0151(SD) bits for cohort M1, with the largest value being 0.06 bits, indicating that despite non-zero sociability, the precise mouse position at any single moment is difficult to predict from the network.

Revised version: Another possible measure of sociability is the mutual information between a single mouse’s location within the territory and the location of the rest of the cohort, which tells us how accurately the position of a single mouse can be predicted if the positions of all other mice are known (see details in). The possible values of the mutual information is between 0 and 2 bits, where 0 bits means no predictability, and 2 bits means perfect predictability. In our Eco-HAB data, the average mutual information for each mouse is 0.0323 ± 0.0151(SD) bits for cohort M1, with the largest value being 0.06 bits, indicating that despite non-zero sociability, the precise mouse position at any single moment is difficult to predict from the network.

Effect of impairing neuronal plasticity in the PL on subterritory preferences and sociability

Earlier version: As a next step we investigate the effects of impairing neuronal plasticity in the prelimbic cortex (PL), the brain structure containing neural circuits indispensable for both maintaining proper social interactions and encoding individual preferences [22]. To that end, we inject animals with a Tissue Inhibitor of MetalloProteinases (TIMP-1), an enzyme regulating the activity of synaptic plasticity proteins. Changing its physiological levels was previously shown to disrupt the neuronal plasticity in various brain structures [43, 44], including the prefrontal cortex (PFC) [33, 45]. Now, we ask if we can identify changes in behavioral patterns after the mice have been injected with nanoparticles (NP) gradually releasing TIMP-1 (NP-TIMP-1) into the PL [37]. It has been demonstrated in the Eco-HAB that it can reduce the mice’s interest in chasing other animals (a proxy for their social ranks), and diminish persistence in seeking reward related to social olfactory cues [45]. Here, we measure the behavior of a cohort of N = 15 C57BL6/J mice, before and after the injection of NP-TIMP-1. A cohort of mice is introduced into the Eco-HAB and their free behavior is measured for 10 days (see schematics in Fig. 3A). Then, neuronal activity in the PL of the subjects is impaired by injecting nanoparticles releasing TIMP-1 into the PL. After recovery animals are re-introduced into the Eco-HAB, and their behavior is measured for another 10 days. As a control we also have a cohort (male cohort M4, N = 9) that is injected with nanoparticles loaded with bovine serum albumin, a physiologically neutral substance having no impact on neuronal plasticity (BSA, vehicle). The detailed experimental procedure can be found in [45]. To provide a perspective on both sexes, a female cohort is also included in this study (female cohort F1, N = 13); it was processed identically to the experimental group of males described above. For each five days of the experiment, we infer a pairwise model (eq. 2) and study the changes of the inferred interactions, as well as individual preferences for specific spaces within the territory. The choice of five-day aggregated data for analysis is in line both with the proper timescales needed for the pairwise maximum entropy model to not overfit, and with the literature that TIMP-1 release from the TIMP-1-loaded nanoparticles is stable for 7-10 days after injection [37] (i.e. 2-5 days after the mice are reintroduced to Eco-HAB).

Revised version: As a next step we investigate the effects of impairing neuronal plasticity in the prelimbic cortex (PL), the brain structure containing neural circuits indispensable for both maintaining proper social interactions and encoding individual preferences [24]. To that end, we inject animals with a Tissue Inhibitor of MetalloProteinases (TIMP-1), an enzyme regulating the activity of synaptic plasticity proteins. Changing its physiological levels was previously shown to disrupt the neuronal plasticity in various brain structures [43, 44], including the prefrontal cortex (PFC) [35] where an overexpression of TIMP-1 is found to block the activity of matrix metalloproteinases and prevent the induction of late long-term potentiation in vivo. More specifically to the PL, it has been recently demonstrated in the Eco-HAB that injecting nanoparticles (NP) gradually releasing TIMP-1 (NP-TIMP-1) [39] in the PL can reduce the mice’s interest in chasing other animals (a proxy for their social ranks), and diminish persistence in seeking reward related to social olfactory cues [45] – which supports the idea that TIMP-1 has regional-specific effects on behavioral processes.

Added in the revised version: Here, we measure the behavior of a cohort of N = 15 C57BL/6J male mice, before and after the injection of NP-TIMP-1. A cohort of mice is introduced into the Eco-HAB and their free behavior is measured for 10 days (see schematics in Fig. 3A). Then, neuronal activity in the PL of the subjects is impaired by injecting nanoparticles releasing TIMP-1 into the PL. After recovery animals are re-introduced into the Eco-HAB, and their behavior is measured for another 10 days. As a control we also have a cohort (male cohort M4, N = 9) that is injected with nanoparticles loaded with bovine serum albumin, a physiologically neutral substance having no impact on neuronal plasticity (BSA, vehicle). The detailed experimental procedure can be found in [45]. To provide a perspective on both sexes, a female cohort is also included in this study (female cohort F1, N = 13); it was processed identically to the experimental group of males described above. For each five days of the experiment, we infer a pairwise model (Eq. 2) and study the changes of the inferred interactions, as well as individual preferences for specific spaces within the territory. The choice of fiveday aggregated data for analysis is in line both with the proper timescales needed for the pairwise maximum entropy model to not overfit, and with the literature that TIMP-1 release from the TIMP-1-loaded nanoparticles is stable for 7-10 days after injection [39] (i.e. 2-5 days after the mice are reintroduced to Eco-HAB).

Quantification of sociability, and the impact of the impaired neuronal plasticity in the prelimbic cortex (PL).

Earlier version: (A) The schematic of the experiment, in which neuronal plasticity in the PL of the tested subjects was impaired with TIMP-1 treatment. A cohort of C57BL6/J male mice (N = 15) was tested in Eco-HAB for 10 days, and then removed from the cages for neuronal plasticity manipulation procedures. After a recovery period, they were placed back in Eco-HAB for another 10 days. For each of the five-day aggregate of the experiment, both before and after TIMP-1 treatment, we plot (B) the model-inferred interactions Jij, (C) preference for food compartments Δhi, (D) conditional log-likelihood for the pairwise model, l(1,2, the independent model, l(1), and the baseline null model, l(0), (E) mutual information between single mouse position and the rest of the network given by the inferred pairwise model.

Revised version: (A) The schematic of the experiment, in which neuronal plasticity in the PL of the tested subjects was impaired with TIMP-1 treatment. A cohort of C57BL/6J male mice (N = 15) was tested in Eco-HAB for 10 days, and then removed from the cages for neuronal plasticity manipulation procedures. After a recovery period, they were placed back in Eco-HAB for another 10 days. For each of the five-day aggregate of the experiment, both before and after TIMP-1 treatment, we plot (B) the model-inferred interactions Jij, (C) preference for food compartments Δhi, (D) conditional log-likelihood for the pairwise model, l(1,2, the independent model, l(1), and the baseline null model, l(0), (E) mutual information between single mouse position and the rest of the network given by the inferred pairwise model. (F-J) Same as (A-E), now for male C57BL/6J mice subject to injection of BSA-infused nanoparticles, a control which does not impair neuronal plasticity in the prelimbic cortex (PL) (cohort M4, N = 9).

Earlier version: We can assess the change in both the interaction strength Jij and the individual preferences for the compartments containing food Δhi following the prolonged release of TIMP-1. The individual preferences for compartment containing food shows an increase in both its mean and its variance across all mice following treatment (Fig. 3C), with a return to pre-treatment levels after five days, consistent with the time course of TIMP-1 release [37]. In comparison, the control cohort M4 shows the increase in preference for the compartments containing food after injection of the BSA-loaded nanoparticles, which does not return to base level after 5 days (SI Fig. S10C). For the interaction strength, as shown by Fig. 3B, the variability of interactions Jij is not significantly different before and after the TIMP-1 injection for the male cohort M1 (two-samples F -test for equal variance with Bonferroni correction; see Materials and Methods), although there exists a few outliers with strong interactions, which again returns to base level after 5 days of TIMP-1 injection. The increase of interaction variability is significant in the female cohort, when comparing the first 5 days after treatment with the first 5 days before treatment (SI Fig. S11), and is not observed for the male BSA cohort M4 (SI Fig. S10). We can also ask how TIMP-1 induced modification of PL plasticity affects individual mice, by comparing the pairwise specific interactions Jij before and after drug treatment. However, we cannot conclude much as Pearson’s correlation coefficient between Jij shows almost no significant correlation across the four time periods in the above datasets for both the TIMP-1 treated cohort and the BSA-treated control cohort (SI Fig. S12).

Revised version: We can assess the change in both the interaction strength Jij and the individual preferences for the compartments containing food Δhi following the prolonged release of TIMP-1. The individual preferences for compartment containing food shows an increase in both its mean and its variance across all mice following treatment (Fig. 3C), with a return to pre-treatment levels after five days, consistent with the time course of TIMP-1 release [39]. In comparison, the control cohort M4 shows the increase in preference for the compartments containing food after injection of the BSA-loaded nanoparticles, which does not return to base level after 5 days (Fig. 3H). For the interaction strength, as shown by Fig. 3B, the variability of interactions Jij is not significantly different before and after the TIMP-1 injection for the male cohort M1 (two-samples F-test for equal variance with Bonferroni correction; see), although there exists a few outliers with strong interactions, which again returns to base level after 5 days of TIMP-1 injection. The increase of interaction variability is significant in the female co-hort, when comparing the first 5 days after treatment with the first 5 days before treatment (SI Fig. S10), and is not observed for the male BSA cohort M4 (Fig. 3G). We can also ask how TIMP-1 induced modification of PL plasticity affects individual mice, by comparing the pairwise specific interactions Jij before and after drug treatment. However, we cannot conclude much as Pearson’s correlation coefficient between Jij shows almost no significant correlation across the four time periods in the above datasets for both the TIMP-1 treated cohort and the BSA-treated control cohort (SI Fig. S11).

Earlier version: To quantify the sociability of the entire cohort, we compute conditional likelihoods as introduced in the previous paragraphs, as it measures how much the pairwise model explains the observed data compared to a model where mice behave independently. Fig. 3D shows that for cohort M1, the model’s likelihoods sharply increase following treatment, meaning that the behavior is more predictable. Represented by the independent model, the individual compartment preferences explain most of this increase, suggesting that TIMP-1 treatment reorganizes preferences for specific subterritories. These differences decay back to pre-treatment levels after 5 days, following the time course of drug release. A slightly smaller increase in model’s likelihood is observed in the control cohort M4 after BSA injection (SI Fig. S10D), suggesting that at least part of the change in compartment preferences can be due to the injection procedure rather than change in the neuronal plasticity itself. In contrast, the increasing model likelihood is not found in the female cohort F1, where the conditional likelihood remains constant after TIMP-1 treatment. However, the contribution of the pairwise interaction is increased (SI Fig. S11E), which points to a sex specificity of observed effects.

Revised version: To quantify the sociability of the entire cohort, we compute conditional likelihoods as introduced in the previous paragraphs, as it measures how much the pairwise model explains the observed data compared to a model where mice behave independently. Fig. 3D shows that for cohort M1, the model’s likelihoods sharply increase following treatment, meaning that the behavior is more predictable. Represented by the independent model, the individual compartment preferences explain most of this increase, suggesting that TIMP-1 treatment reorganizes preferences for specific subterritories. These differences decay back to pre-treatment levels after 5 days, following the time course of drug release. A slightly smaller increase in model’s likelihood is observed in the control cohort M4 after BSA injection (Fig. 3I), suggesting that at least part of the change in compartment preferences can be due to the injection procedure rather than change in the neuronal plasticity itself. In contrast, the increasing model likelihood is not found in the female cohort F1, where the conditional likelihood remains constant after TIMP-1 treatment. However, the contribution of the pairwise interaction is increased (SI Fig. S10E), which points to a sex specificity of observed effects.

Earlier version: This observation is further confirmed by the sociability measure of mutual information between single mouse location and the positions of the rest of the cohort, which was introduced in previous paragraphs. The mutual information either does not change (for the male cohort M1, Fig. 3E), or increases (for the female cohort F1, SI Fig. S11) after the injection of TIMP-1.

Revised version: This observation is further confirmed by the sociability measure of mutual information between single mouse location and the positions of the rest of the cohort, which was introduced in previous paragraphs. The mutual information either does not change (for the male cohort M1, Fig. 3E, or increases (for the female cohort F1, SI Fig. S10) after the injection of TIMP-1.

Impaired neuronal plasticity in the PL affects the structure of social interactions

Earlier version: The increasing variability of pairwise interactions and the non-decreasing mutual information between single mouse location and the location of the rest of the group upon TIMP-1 requires further investigation in the face of previous results showing that injecting TIMP-1 to the PL of wide type animals reduces their sociability. Thus, we examined the detailed group structure of pairwise interactions. We define the dissatisfaction triplet index (DTI) among a triplet of mice as Fijk ≡ − JijJjkJki if and only if among the three pairwise interactions among mice i, j and k, exactly one of them is negative (see Fig. 4A for schematics), and otherwise zero. Notice that DTI is analogous to the concept of “frustration” in physics of disordered systems. A positive DTI means a triplet of pairwise interactions where all of them cannot be satisfied simultaneously – for instance, if mouse i likes to be with j and k, but j and k do not like to be together. We define the global DTI by averaging the local DTI’s across all possible triplets of mice. The larger the global DTI is for a cohort, the more difficult it is for the cohort to form cliques with multiple mice where the interactions among each possible pairs are positive, which may suggest possible difficulty in transmitting information between mice.

Revised version: The increasing variability of pairwise interactions and the non-decreasing mutual information between single mouse location and the location of the rest of the group upon TIMP-1 requires further investigation in the face of previous results showing that injecting TIMP-1 to the PL of wide type animals reduces their sociability. Thus, we examined the detailed group structure of pairwise interactions. We define the dissatisfaction triplet index (DTI) among a triplet of mice as Fijk ≡ − JijJjkJki if and only if among the three pairwise interactions among mice i, j and k, exactly one of them is negative (see Fig. 4A for schematics), and otherwise zero. Notice that DTI is analogous to the concept of “frustration” in physics of disordered systems. A positive DTI means a triplet of pairwise interactions where all of them cannot be satisfied simultaneously – for instance, if mouse i likes to be with j and k, but j and k do not like to be together. We define the global DTI by averaging the local DTI’s across all possible triplets of mice. The larger the global DTI is for a cohort, the more difficult it is for the cohort to form cliques with multiple mice where the interactions among each possible pairs are positive, which may suggest possible difficulty in transmitting information between mice. As shown in Fig. 4B, PL-targeted plasticity disruption with TIMP-1 significantly increases the global DTI for the male mice cohort M1 and the female cohort F1 (two-sample Welch’s t-test, variability from random halves of the data; see for details of the significant test). In contrast, in the control cohort M4, injecting male mice with BSA either does not significantly change the global DTI or decreases it. Notably, the difference of the DTI is not due to the control group M4 has less mice, as sub-sampling both on the level of the inferred interactions (SI Fig. S12) and on the level of the mice locations (SI Fig. S13) give the same DTI for cohorts M1 and F1. This increase of the global DTI is due to the increasing variance of the interaction Jij, which is related to more of the negative interactions. Randomly shuffling Jij does not change the global DTI, indicating that no network structure was found that contributes to this global DTI (SI Fig. S14).

Effect of TIMP-1 on the structure of the interaction network.

Earlier version: (A) Schematics of how triplets of mice may enter a state of “dissatisfaction” due to competitive pairwise interactions. Dissatisfaction reduces the space of preferable states due to competitive interactions. (B) The global dissatisfaction triplet index (DTI), F, computed using inferred interaction from 5-day segments of the data shows that for both male and female mice treated with TIMP-1, the global DTI is significantly increased after drug treatment. Two-sided Welch’s t -test is performed to test the significance for the difference of the global DTI between the first 5 days after drug injection against the other 5-day segments of the data. Error bars estimate the data variability, which is generated by taking random halves of the data (see Materials and Methods for details).

Revised version: (A) Schematics of how triplets of mice may enter a state of “dissatisfaction” due to competitive pairwise interactions. Dissatisfaction reduces the space of preferable states due to competitive interactions. (B) The global dis-satisfaction triplet index (DTI), F, computed using inferred interaction from 5-day segments of the data shows that for both male and female mice treated with TIMP-1, the global DTI is significantly increased after drug treatment. Two-sided Welch’s t-test is performed to test the significance for the difference of the global DTI between the first 5 days after drug injection against the other 5-day segments of the data. Error bars estimate the data variability, which is generated by taking random halves of the data (see for details).

Removed in the revised version: As shown in Fig. 4B, PL-targeted plasticity disruption with TIMP-1 significantly increases the global DTI for the male mice cohort M1 and the female cohort F1 (two-sample Welch’s t-test, variability from random halves of the data; see Materials and Methods for details of the significant test). In contrast, in the control cohort M4, injecting male mice with BSA either does not significantly change the global DTI or decreases it. Notably, the difference of the DTI is not due to the control group M4 has less mice, as subsampling both on the level of the inferred interactions (SI Fig. S13) and on the level of the mice locations (SI Fig. S14) give the same DTI for cohorts M1 and F1. This increase of the global DTI is due to the increasing variance of the interaction Jij, which is related to more of the negative interactions. Randomly shuffling Jij does not change the global DTI, indicating that no network structure was found that contributes to this global DTI (SI Fig. S15).

Discussion

Earlier version: We demonstrate how the joint probability distribution of the mice positions in the Eco-HAB can be used to quantify sociability. By building a pairwise interaction model whose parameters are learned directly from the data, we quantify how much the combined activities of all mice in the cohort are influenced by their individual preferences and how much by the social context. This approach shows that pairwise interactions between mice are sufficient to describe the statistics of collective behavior in larger groups. Additionally, the pairwise interaction model can capture changes in the social interactions of the network induced by alterations in the neuronal plasticity of the prelimbic cortex (PL) in the tested subjects. The Eco-HAB, combined with this analytical approach, provides a toolbox to quantify sociability in mice, which can be applied generally to different mouse strains to study various behavioral phenotypes, including characteristics associated with neurodevelopmental disorders such as autism. Compared to traditional experimental methods like the three-chamber test, our study combines the advantages of an ecologically relevant and automatic experimental apparatus with the powerful tools of statistical inference. This disentangles the effects of individual preferences versus pairwise social interactions in generating the patterns of mouse positions within their territory.

Revised version: We demonstrate how the joint probability distribution of the mice positions in the Eco-HAB can be used to quantify sociability. By building a pairwise interaction model whose parameters are learned directly from the data, we quantify how much the combined activities of all mice in the cohort are influenced by their individual preferences and how much by the social context. This approach shows that, within the setup of the seminatu-ralistic Eco-HAB experiments, pairwise interactions between mice are sufficient to describe the statistics of collective behavior in larger groups. Additionally, the pair-wise interaction model can capture changes in the social interactions of the network induced by alterations in the neuronal plasticity of the prelimbic cortex (PL) in the tested subjects. The Eco-HAB, combined with this analytical approach, provides a toolbox to quantify sociability in mice, which can be applied generally to different mouse strains to study various behavioral phenotypes, including characteristics associated with neurodevelopmental disorders such as autism. Compared to traditional experimental methods like the three-chamber test, our study combines the advantages of an ecologically relevant and automatic experimental apparatus with the powerful tools of statistical inference. The use of statistical inference methods such as maximum entropy models disentangles the effects of individual preferences versus pairwise social interactions in generating the patterns of mouse positions within their territory.

The challenge in studying social behavior lies in finding a balance between being specific enough to capture the properties of sociability while avoiding the loss of generalizability. Including excessive details, such as the classification of precise social behavior among mice, may lead to a more accurate description of the specific mice cohort studied, such as the construction of a precise social network. However, it is difficult to assess comparability across different cohorts of mice. Alternatively, as used in this paper, one can construct minimal models and use the ensemble statistics of the models to quantify social properties of a mouse strain without explicitly constructing social networks for each cohort. For example, our study found that the inferred interaction has similar ensemble statistics across four different male cohorts of the same strain but differs across different sexes. This provides evidence to support our argument for a coarse-grained description of mouse social behavior.

Another challenge in studying social behavior lies in the interplay of timescales. We found in this paper that in order to gather enough statistics and to avoid over-fitting, we need to construct probabilistic models using aggregated five-day data, which poses a challenge to balance model construction with enough data versus studying the temporal evolution of sociability. Is the variability of the social network a true property of the social interaction of the mice cohort, or is it due to variabilities of the inferred model caused by the finite data? To address this question, one needs to consider various timescales. For example, mice-mice interactions occur at a much shorter timescale compared to the timescale of changes in the social network, while in between, there are the timescales of adaptation to the new environment and the circadian cycle. These issues need to be addressed using a combination of theoretical tools and experimental validation methods in future works.

Additionally, we have simplified our analysis by focusing on a 6-hour time window each day, during which the rate of locomotor activity is most stable. This approach allowed us to circumvent issues related to individual or strain differences in the circadian cycle, such as the observed “lunch hour” in C57 male mice. One avenue for future research involves reintroducing the circadian cycle as a latent variable to better explain the system. However, caution must be exercised to differentiate between group behavior influenced by the circadian cycle of individual mice and emergent behavior resulting from interactions.

Added in the revised version: While not the main focus of this paper, we tested our methods of quantifying sociability on both female and male mice of the same strain undergoing the same alteration of neuronal plasticity in the PL. Interestingly, while both male and female mice demonstrate an increase of competitive pairwise interactions upon such alteration, they also exhibit many differences in sociability: for example, the individual preferences increase more in male mice after TIMP-1 injection than the female mice.

Earlier version: We will now discuss the relationship between our study and that of Shemesh et al., wherein the authors applied a similar approach, investigating the social behavior of groups of 4 mice in a complex experimental environment using statistical modeling of the joint probability distribution of mice locations [39]. They found that triplet interactions are necessary to describe collective behavior, while we found that triplet interactions can be predicted by the pairwise model. We suspect the difference in our results could arise from three factors. First, the experimental arena is quite different between the Eco-HAB and what was used in Shemesh et al. – while the Eco-HAB mimics a naturalistic environment with tubes and connected chambers, the experimental arena in Shemesh et al. contains only one chamber, which could cause a difference in how mice interact. Second, even if the interaction patterns do not change, the nature of the data is different: our data is more coarse-grained spatially, as the state of each mouse is determined by the large compartment it is in, whereas in Shemesh et al., the location is more precise (e.g., a door, a pillar, etc.). As suggested by a comparison to renormalization theory in physics, at coarser spatial scales, the importance of higher-order interactions is likely to decrease. Finally, our studies include more mice (10 to 15) compared to Shemesh et al. (4 mice), which may also influence the importance of higher-order interactions. Larger group size also means that including triplet interactions in our model causes overfitting, which restricts the models we have access to to pairwise interactions. To further investigate these effects, future experiments in Eco-HAB could include mice cohorts of smaller sizes.

Revised version: We will now discuss the relationship between our study and that of Shemesh et al., wherein the authors applied a similar approach, investigating the social behavior of groups of 4 mice in a complex experimental environment using statistical modeling of the joint probability distribution of mice locations [21]. Both Shemesh et al. and us uses maximum entropy models to analyze co-localization patterns of a group of mice housed in complex environments. While Shemesh et al. found that triplet interactions are necessary to describe collective behavior, we found that triplet interactions can be predicted by the pairwise model. We suspect the difference in our results could arise from three factors. First, the experimental arena is quite different between the Eco-HAB and what was used in Shemesh et al. – while the Eco-HAB mimics a naturalistic environment with tubes and connected chambers, the experimental arena in Shemesh et al. contains only one chamber, which could cause a difference in how mice interact. Second, even if the interaction patterns do not change, the nature of the data is different: our data is more coarse-grained spatially, as the state of each mouse is determined by the large compartment it is in, whereas in Shemesh et al., the location is more precise (e.g., a door, a pillar, etc.). As suggested by a comparison to renormalization theory in physics, at coarser spatial scales, the importance of higher-order interactions is likely to decrease. Finally, our studies include more mice (10 to 15) compared to Shemesh et al. (4 mice), which may also influence the importance of higher-order interactions. Larger group size also means that including triplet interactions in our model causes overfitting, which restricts the models we have access to to pairwise interactions. To further investigate these effects, future experiments in Eco-HAB could include mice cohorts of smaller sizes. More generally, this discrepancy when looking at different choices of variables raises the issue that when studying social behavior of animals in a group, it is important to test and compare interaction models with different complexity (e.g. pairwise or with higher-order interactions). Furthermore, since pair-wise maximum entropy model is one of the simplest of all maximum entropy models that can describe interactions among individuals, it serves as an excellent starting point to describe collective and social behavior in animals.

How do we move forward, and what is the ideal experiment to study social behaviors? We believe that Eco-HAB offers a balance between a semi-natural environment and controllability, which works well in studying social behavior. One direction for future experimental studies is to focus on the biological function of social interactions. For example, how do mouse cohorts respond to novel odors and transmit information among the cohort? What is the speed of information transmission related to sociability? The current configuration of the Eco-HAB already allows for the introduction of novel odors accessible to all mice, while the next generation of experiments will localize the introduction of information to individuals. From the analysis perspective, as presented in this manuscript, our current model is purely static. Our model describes the joint probability distribution of mice positions within the territory at concurrent time points and does not model the dynamics of the cohort. To take into account the dynamic aspect of social behaviors, such as dominant mice actively chasing others, one will need to build dynamical models of interaction. For example, this can be done by modeling the probability of transitioning to another compartment of each mouse as a function of the history of its previous location and the locations of all other mice [46].

Materials and methods

Animals

Removed in the revised version: Animals were treated in accordance with the ethical standards of the European Union (directive no. 2010/63/UE) and Polish regulations. All the experiments were pre-approved by the Local Ethics Committee no. 1 in Warsaw, Poland. C57BL/6 male mice were bred in the Animal House of Nencki Institute of Experimental Biology, Polish Academy of Sciences or Mossakowski Medical Research Centre, Polish Academy of Sciences.

Earlier version: The animals entered the experiments when 2-3 month old. They were littermates derived from several breeding pairs. The mice were transferred to the animal room at least 2 weeks before the experiments started and put in the groups of 12-15 in one cage (56 cm× 34 cm × 20cm) with enriched environment (tubes, shelters, nesting materials). They were kept under 12h/12h light-dark cycle. The cages were cleaned once per week. Five cohorts of mice were used in the analysis of this manuscript (see details in Table 1).

Revised version: Animals were treated in accordance with the ethical standards of the European Union (directive no. 2010/63/UE) and Polish regulations. All the experiments were pre-approved by the Local Ethics Committee no. 1 in Warsaw, Poland. C57BL/6J male and female mice were bred in the Animal House of Nencki Institute of Experimental Biology, Polish Academy of Sciences or Mossakowski Medical Research Centre, Polish Academy of Sciences. The animals entered the experiments when 2-3 month old. They were littermates derived from several breeding pairs. The mice were transferred to the animal room at least 2 weeks before the experiments started and put in the groups of 12-15 in one cage (56 cm× 34 cm× 20cm) with enriched environment (tubes, shelters, nesting materials). They were kept under 12h/12h light-dark cycle. The cages were cleaned once per week. Five cohorts of mice were used in the analysis of this manuscript (see details in Table 1).

Exclude inactive and dead mice from analysis

Mouse whose trajectory does not cover all four compartments within the 6-hour period for at least one day of the experiment is defined as inactive, and excluded from the analysis. Including inactive mice in the maximum entropy model will results in unstable learned parameters, as shown by bootstrapped results. For the same mouse cohort before and after injection of drug (M1, M4, and F1), if a mouse is dead or inactive in either phase of the experiment, its trajectory is masked out from the data for consistency of comparison before and after. Specifically, for cohort F1, mouse number 13 (in the original ordering of the 14 mice) is inactive after the drug application. For cohort M4, mouse number 3 and 11 (in the original ordering of the 12 mice) died after surgery, mouse number 9 (in the original ordering) is inactive in the 10th day after drug injection. The total number of mice used in the analysis is given in Table 1.

Longitudinal observation of social structure in the Eco-HAB

Cohorts of mice with the same gender and same strain were placed in the Eco-HAB systems and observed for 10 days, removed from the system to undergo stereotaxic injections with TIMP-1 loaded nanoparticles. After 4 to 6 days of recovery, the mice were placed back to a cleaned Eco-HAB, and observed for 10 days.

Activity level

The activity level for a given mouse i during a given time period (ti, tf) on day d, is computed by counting the number of times the mouse passes by any antenna, and denoted by .

Earlier version: Averaging this quantity over all N mice, one obtain the mean activity level for all mice during a given time period. Mathematically,.

Revised version: Averaging this quantity over all N mice, one obtain the mean activity level for all mice during a given time period. Mathematically, . The standard deviation across all days is the day-to-day variability of mean activity level.

Removed in the revised version: The standard deviation across all days is the day-to-day variability of mean activity level.

Averaging this quantity over all T days, one obtain the mean activity level for each mouse. Mathematically,. The standard deviation across all mice is the mouse-to-mouse variability of the mean activity level.

Mice location

The raw data consists of time points when mice cross an antenna, as well as the identity of the specific antenna, which are placed at the ends of the four tunnels. The location of a mouse at any given time point is deduced from the most recent time stamps before and after the current time point. For simplicity, for the time points when a mouse is in the tunnel, the location of the mouse is set to be the compartment it will enter. The time resolution is set to 2 seconds, as two adjacent time stamps with separation less than 2 seconds are likely an artifact of mice sniffing the tunnel and returning to the previous compartment.

Pseudocounts

Observing the mice for a finite amount of time means sometimes we have the situation where the mice is stuck in the same compartment for the entire 6 hours of observation. This is not common, but this messes up our statistical inference or model building procedure. To avoid this situation, we use pseudocounts that smoothen the observed statistics. We define and where q = 4 is the number of possible states, T is the total number of time points in the data, and λ is the parameter for the pseudocount. In our analysis, after scanning through a range of values for λ, we set λ = 8 around which value the outcome remained largely unchanged.

The probability model

Removed in the revised version: Gauge fixing for the local field

Earlier version: the probability model is equivalent for the local fields upon a constant, i.e. P (h) and P (hir + δhi) are equivalent. We overcome this redundancy by enforcing the sum of all local fields for each mouse to be zero, i.e. ∑r hir = 0.

Revised version: Gauge fixing for the local field : the probability model is equivalent for the local fields upon a constant, i.e. P (h) and P (hir + δhi) are equivalent. We overcome this redundancy by enforcing the sum of all local fields for each mouse to be zero, i.e. ∑ r hir = 0.

Learning the probability model

Earlier version: We train the model using gradient descent, at each learning step k updating the parameters by and , where α = 0.25 ∼0.8 is the step size of learning. The stopping condition is set such that when the difference between the model predicted correlation and magnetization is less than the data variability, estimated by extrapolation from random halves of the data. In addition, because we are interested in quantifying social properties of mice cohort using the statistics of learned parameters, we add to the stopping condition that the mean and the variation of inferred interaction reach a stable value, with change less than 0.005 over 100 learning steps.

Revised version: We train the model using gradient descent, at each learning step k updating the parameters by and , where α = 0.25 ~ 0.8 is the step size of learning. The stopping condition is set such that when the difference between the model predicted correlation and magnetization is less than the data variability, estimated by extrapolation from random halves of the data. In addition, because we are interested in quantifying social properties of mice cohort using the statistics of learned parameters, we add to the stopping condition that the mean and the variation of inferred interaction reach a stable value, with change less than over 100 learning steps.

Computing higher-order correlations

Earlier version: The connected three-point correlation function gives the frequency of finding three mice in the same compartment, subtracting the contributions from the mean and the pair-wise correlation. Mathematically, If we only subtract the individual preference, then we define

Revised version: The connected three-point correlation function gives the frequency of finding three mice in the same compartment, subtracting the contributions from the mean and the pair-wise correlation. Mathematically,

Added in the revised version: If we only subtract the individual preference, then we define

Comparing in-state probability between model prediction and data

Given time-series data and the inferred joint probability distribution of mice location, we can compare the in-state probability of single mouse, as given by model prediction versus data observation.

Earlier version: More precisely, for each time point t, given the inferred compartment preference hir for mouse i, the inferred pairwise interaction Jij, and the position of all other mice j ≠ i, we use the pairwise maximum entropy model (eq. 2) to compute the marginal probability of mouse i being in each of the four compartments as the model-predicted “in-state probability”, . These model-predicted in-state probabilities are then binned according to their percentiles across all observed time points for each mouse and each box, such that the number of time points in each bin is equal. Then, for all time points that belong to each bin, the frequency count of whether mouse i is actually observed in compartment r is computed as the observed in-state probability. Agreement between the model-predicted and the observed in state probability across all the bins shows the model is an unbiased estimator, which is what we see after averaging - the in-state probability across the four compartments for fixed percentiles in SI Fig. S5.

Revised version: More precisely, for each time point t, given the inferred compartment preference hir for mouse i, the inferred pairwise interaction Jij, and the position of all other mice j ≠ i, we use the pairwise maximum entropy model (Eq. 2) to compute the marginal probability of mouse i being in each of the four compartments as the model-predicted “in-state probability”, . These model-predicted in-state probabilities are then binned according to their percentiles across all observed time points for each mouse and each box, such that the number of time points in each bin is equal. Then, for all time points that belong to each bin, the frequency count of whether mouse i is actually observed in compartment r is computed as the observed in-state probability. Agreement between the model-predicted and the observed instate probability across all the bins shows the model is an unbiased estimator, which is what we see after averaging the in-state probability across the four compartments for fixed percentiles in SI Fig. S5.

Compute mutual information between single mouse position and the rest of the network

Earlier version: The mutual information between single mouse position and the rest of the network is a measure of collectiveness. For Eco-HAB with 4 compartments, the mutual information is between 0 and 2 bits. If the mutual information is close to 2 bits, knowing where other mice are is a perfect predictor for the position of single mouse. If the mutual information is close to 0 bits, knowing where other mice are do not help predicting the position of the singled-out mouse. The mutual information can be computed as the difference of the entropy of mouse i and the conditional entropy of mouse i with respect to the state of all other mice. Mathematically, where the entropy of mouse i is and the conditional entropy is computed using the conditional probability given {σj} and the inferred pair-wise data, and averaged over all observed data patterns {σj(t)}. To reach the final results, we approximate the ensemble average over all possible mice configurations with a temporal average over all observed mice configuration in the data. We also replace the true underlying conditional probability of P (σi|{σj}j≠i) with the inferred pairwise probability model P (2).

Revised version: The mutual information between single mouse position and the rest of the network is a measure of collectiveness. For Eco-HAB with 4 compartments, the mutual information is between 0 and 2 bits. If the mutual information is close to 2 bits, knowing where other mice are is a perfect predictor for the position of single mouse. If the mutual information is close to 0 bits, knowing where other mice are do not help predicting the position of the singled-out mouse. The mutual information can be computed as the difference of the entropy of mouse i and the conditional entropy of mouse i with respect to the state of all other mice. Mathematically, where the entropy of mouse i is and the conditional entropy is computed using the conditional probability given {σj} and the inferred pair-wise data, and averaged over all observed data patterns {σj(t)}.

Added in the revised version: To reach the final results, we approximate the ensemble average over all possible mice configurations with a temporal average over all observed mice configuration in the data. We also replace the true underlying conditional probability of P (σi|{σj}j≠i) with the inferred pairwise probability model P (2).

Generate errorbars using random bootstrapped halves of the data

Earlier version: Error bars of the observed statistics 𝒪 (e.g. pairwise correlation, Cij, and probability in each compartment, mir), the inferred parameters 𝒫 (e.g. pairwise interaction Jij and compartment preference hir), and the sub-sequent results ℛ(e.g. the entropy, S(1,2) and S(1), and the dissatisfaction triplet index F) are bootstrap errors generated by repeatedly taking random halves of the data and computing the deviations in the mean. Specifically, each data set (at least 6 hours in duration) is first divided into time bins of 400 seconds. The length of the time bin is chosen such that it is longer than twice the correlation time for each mouse. Then, random halves of the time bins are chosen to compute the observables, as well as used to train a specific pairwise maximum entropy model, which generates a specific set of learned parameters. The deviation across the random halves, σbs, can be extrapolated to the full dataset by .

Revised version: Error bars of the observed statistics 𝒪 (e.g. pairwise correlation, Cij, and probability in each compartment, mir), the inferred parameters 𝒫 (e.g. pairwise interaction Jij and compartment preference hir), and the subsequent results ℛ (e.g. the entropy, S(1,2) and S(1), and the dissatisfaction triplet index F) are bootstrap errors generated by repeatedly taking random halves of the data and computing the deviations in the mean. Specifically, each data set (at least 6 hours in duration) is first divided into time bins of 400 seconds. The length of the time bin is chosen such that it is longer than twice the correlation time for each mouse. Then, random halves of the time bins are chosen to compute the observables, as well as used to train a specific pairwise maximum entropy model, which generates a specific set of learned parameters. The deviation across the random halves, σbs, can be extrapolated to the full dataset by.

Test of significance for comparing observables and inferred parameters

Earlier version: To perform significance test across different days of the experiment, we used the Welch’s t -test for the mean of the inferred interaction ⟨Jij ⟩, the mean of the food preference ⟨Δhi ⟩, the mean of the mutual information between single mouse position and the rest of the network given by the inferred pairwise model, and for the global dissatisfaction triplet index (DTI), F. We used two-sample F -test for the variability of the inferred interaction. Because we are comparing between pairs of the 5-day aggregate data – the first 5 days before drug injection, the last 5 days before drug injection, the first 5 days after drug injection, and the last 5 days after drug injection – we conduct significance tests using Bonferroni corrections. The number of tests performed for such pairwise comparison is 6. In Fig. 3, the asterisks encode the following p-values: *, p ≤ 0.05/6; * *, p ≤ 0.01/6; * * *, p ≤ 0.001/6. For the significance test comparing the global DTI, random halves of the 5-day aggregated data is chosen 10 times, each used to learn the interaction parameters and compute the global DTI. The variance of the global DTI across the 10 random halves is used as variation due to finite amount of data, and is adjusted by . Two-tailed tests are performed, and the Bonferroni correction is applied as the total number of tests performed for the pairwise comparisons for the 5-day aggregated data before and after pharmacological intervention is 6. In Fig. 4, the asterisks encode the following p-values: *, p ≤ 0.025/6; **, p ≤ 0.005/6; * * *, p ≤ 0.0005/6.

Revised version: To perform significance test across different days of the experiment, we used the Welch’s t-test for the mean of the inferred interaction (Jij), the mean of the food preference (Δhi), the mean of the mutual information between single mouse position and the rest of the network given by the inferred pairwise model, and for the global dis-satisfaction triplet index (DTI), F. We used two-sample F-test for the variability of the inferred interaction. Because we are comparing between pairs of the 5-day aggregate data – the first 5 days before drug injection, the last 5 days before drug injection, the first 5 days after drug injection, and the last 5 days after drug injection – we conduct significance tests using Bonferroni corrections. The number of tests performed for such pairwise comparison is 6. In Fig. 3, the asterisks encode the following p-values: ∗, p ≤ 0.05/6; ∗∗, p ≤ 0.01/6; ∗ ∗ ∗, p ≤ 0.001/6.

Added in the revised version: For the significance test comparing the global DTI, random halves of the 5-day aggregated data is chosen 10 times, each used to learn the interaction parameters and compute the global DTI. The variance of the global DTI across the 10 random halves is used as variation due to finite amount of data, and is adjusted by.

Added in the revised version: Two-tailed tests are performed, and the Bonferroni correction is applied as the total number of tests performed for the pairwise comparisons for the 5-day aggregated data before and after pharmacological intervention is 6. In Fig. 4, the asterisks encode the following p-values: ∗, p ≤ 0.025/6; ∗∗, p ≤ 0.005/6; ∗ ∗ ∗, p ≤ 0.0005/6.

Maximum entropy model with triplet interactions

Earlier version: To model the joint probability distribution of mice colocalization pattern, one could in principle increase the number of constrains when constructing the maximum entropy model. One example is the triplet correlation, i.e. the probability of any triplets of mice are being found in the same box,. The corresponding maximum entropy model has the probability distribution where hir is the individual preference of mouse i to be in compartment r and Jij is the pairwise interaction between mouse i and mouse j as before, and Gijk is the triplet interaction among mouse i, mouse j and mouse k. For a cohort with N mice, the numbers of parameters for triplet interactions is N (N− 1)(N− 2)/6. To avoid overfitting, we applied a L2 regularization on the triplet interaction strength, where we now minimize the objective function This is naturally translated to learning steps: , and . The stopping condition is the same as before, focusing on the difference between the model predicted correlation and magnetization being less than the data variability, and that the mean and variation of the inferred pairwise interaction reach a stable value.

Revised version: To model the joint probability distribution of mice colocalization pattern, one could in principle increase the number of constrains when constructing the maximum entropy model. One example is the triplet correlation, i.e. the probability of any triplets of mice are being found in the same box,. The corresponding maximum entropy model has the probability distribution where hir is the individual preference of mouse i to be in compartment r and Jij is the pairwise interaction between mouse i and mouse j as before, and Gijk is the triplet interaction among mouse i, mouse j and mouse k. For a cohort with N mice, the numbers of parameters for triplet interactions is N (N − 1)(N − 2)/6. To avoid overfitting, we applied a L2 regularization on the triplet interaction strength, where we now minimize the objective function

Added in the revised version: This is naturally translated to learning steps: , and. The stopping condition is the same as before, focusing on the difference between the model predicted correlation and magnetization being less than the data variability, and that the mean and variation of the inferred pairwise interaction reach a stable value.

To test the validity of a triplet interaction model, we varied the regularization strength βG. For our largest dataset, the cumulative data consist of all 10 days of the experiment on cohort M1 (N = 15), we performed cross-validation by splitting the 6-hour data into 6 different training-test set combo (5 hours of data in training set, and 1 hour in the test set). As shown by SI Fig. S6, the test set likelihood is found to be maximized when the regularization parameter βG is large, which corresponds to close to zero triplet interactions. This indicates that even for our largest datasets, including the triplet interactions in the model have increased the model complexity beyond what the data allows, and the best way to avoid overfitting is for us to not consider triplet interactions. Hence, the interaction models we choose to study in this manuscript is restricted to pairwise interaction models.